Before you submit this notebook, make sure everything runs as expected in the local test cases. 
Please, paste the solution to the designed cell and do not change anything else.

Also, please, leave your first and last names below

```python
FirstName = "Daniil"
LastName = "Zelezetsky"
```

---

```python
import sys

import numpy as np
import unittest
import time

import collections
import pickle
import io
```

<!-- grade_id: cell-ac8fc52d8a39ccb4, points: 0, grade: false, solution: true -->
```python
import numpy as np

from loss_and_derivatives import LossAndDerivatives

```

### Test 0: Initialization (0.01 points)

<!-- grade_id: cell-3473b7b6ffd64d07, points: 0.01, grade: true, solution: false -->
```python
# do not change this cell
loss_and_derivatives = LossAndDerivatives


```

```python
import json
with open('boston_subset.json', 'r') as iofile:
    dataset = json.load(iofile)
feature_matrix = np.array(dataset['data'])
targets = np.array(dataset['target'])

w = np.array([1., 1.])
x_n, y_n = feature_matrix, targets

# Repeating data to make everything multi-dimentional
w = np.vstack([w[None, :] + 0.27, w[None, :] + 0.22, w[None, :] + 0.45, w[None, :] + 0.1]).T
y_n = np.hstack([y_n[:, None], 2*y_n[:, None], 3*y_n[:, None], 4*y_n[:, None]])
```

### Test 1: MSE derivative (0.24 points)

```python
X_ref = x_n
y_ref = y_n
w_hat = w




reference_mse_derivative = np.array([
    [ 7.32890068, 12.88731311, 18.82128365, 23.97731238],
    [ 9.55674399, 17.05397661, 24.98807528, 32.01723714]
])
reference_l2_reg_derivative = np.array([
    [2.54, 2.44, 2.9 , 2.2 ],
    [2.54, 2.44, 2.9 , 2.2 ]
])
reference_mae_derivative = np.array([
    [0.19708867, 0.19621798, 0.19621798, 0.19572906],
    [0.25574138, 0.25524507, 0.25524507, 0.25406404]
])
reference_l1_reg_derivative = np.array([
    [1., 1., 1., 1.],
    [1., 1., 1., 1.]
])




ref_dict = {'mse_derivative':reference_mse_derivative,'mae_derivative':reference_mae_derivative,
            'l2_reg_derivative':reference_l2_reg_derivative,'l1_reg_derivative':reference_l1_reg_derivative}
```

<!-- grade_id: cell-e3503c286039ec55, points: 0.24, grade: true, solution: false -->
```python
mse_derivative = LossAndDerivatives.mse_derivative(X_ref, y_ref, w_hat)
assert np.allclose(mse_derivative, ref_dict['mse_derivative'], atol=1e-4)
```

### Test 2: MAE derivative (0.25 points)

<!-- grade_id: cell-e2c4124a6f815118, points: 0.25, grade: true, solution: false -->
```python
mae_derivative = LossAndDerivatives.mae_derivative(X_ref, y_ref, w_hat)
assert np.allclose(mae_derivative, ref_dict['mae_derivative'], atol=1e-4)
```

### Test 3: L1 reg derivative (0.15 points)

<!-- grade_id: cell-69473387a23d8dff, points: 0.15, grade: true, solution: false -->
```python
l2_reg_derivative = LossAndDerivatives.l2_reg_derivative(w_hat)
assert np.allclose(l2_reg_derivative, ref_dict['l2_reg_derivative'], atol=1e-4)
```

### Test 4: L1 reg derivative (0.15 points)

<!-- grade_id: cell-3460671ba79fd04a, points: 0.15, grade: true, solution: false -->
```python
l1_reg_derivative = LossAndDerivatives.l1_reg_derivative(w_hat)
assert np.allclose(l1_reg_derivative, ref_dict['l1_reg_derivative'], atol=1e-4)
```

# НАЧИНАЯ СО СЛЕД. ЯЧЕЕК БУДУТ ВЫПАДАТЬ ОШИБКИ ТАК КАК В ДЗ
# НЕ ЗАДАНЫ ПЕРЕМЕННЫЕ mse mae l1reg и l2reg

### Test 5: MSE (0.05 points)

<!-- grade_id: cell-5d56bb4222ee1e7c, points: 0.05, grade: true, solution: false -->
```python
mse = LossAndDerivatives.mse(X_ref, y_ref, w_hat)
assert np.allclose(mse, ref_dict['mse'], atol=1e-4)
```

Output:
```
KeyError                                  Traceback (most recent call last)
Cell In[11], line 2
      1 mse = LossAndDerivatives.mse(X_ref, y_ref, w_hat)
----> 2 assert np.allclose(mse, ref_dict['mse'], atol=1e-4)

KeyError: 'mse'
```

### Test 6: MAE (0.05 points)

<!-- grade_id: cell-8155e619427ac916, points: 0.05, grade: true, solution: false -->
```python
mae = LossAndDerivatives.mae(X_ref, y_ref, w_hat)
assert np.allclose(mae, ref_dict['mae'], atol=1e-4)
```

### Test 7: L2 reg (0.05 points)

<!-- grade_id: cell-c7ccfce79ecbadc3, points: 0.05, grade: true, solution: false -->
```python
l2_reg = LossAndDerivatives.l2_reg(w_hat)
assert np.allclose(l2_reg, ref_dict['l2_reg'], atol=1e-4)
```

### Test 8: L1 reg (0.05 points)

<!-- grade_id: cell-1e1aa415f6d381f7, points: 0.05, grade: true, solution: false -->
```python
l1_reg = LossAndDerivatives.l1_reg(w_hat)
assert np.allclose(l1_reg, ref_dict['l1_reg'], atol=1e-4)   
```